<br>
<br>
<div>
<img  style="float: left; padding-right: 100px; width: 350px" src="../images/africa-no-words.png">
    </div>
    <br>
    <h2 align="center">AI4D-Lab | AI Short Course</h2>
<hr>
<h4 align="center">Zephania Reuben</h4>
<h4 align="center">July, 2022</h4>
<hr>
<h3 align="center">The University of Dodoma - CIVE</h3>
<hr>
<br>
<br>
<h3 align="center">PYTHON FUNDAMENTALS | EXERCISE
<br>

```python
# Initialize Otter
import otter
grader = otter.Notebook("python_exercise_ai4d_lab_2022.ipynb")
```

```python
# Check the IPython version
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

```python
# Import the good stuff
from nose.tools import assert_equal
import numpy as np
import math
import otter
grader = otter.Notebook()
```

**Question 1:**  Write a function `ypoly` to output the value of `y` given `x` and `p`. Function definition is `y = 4 + 5x^p` Ensure that if the `x` is not positive the value of `y` should be `0`

```python
def ypoly(x, p):
    ...
```

```python
grader.check("q1")
```

**Question 2:**  Write a function `perfectly_divisible` to return a list of numbers between 0 and a number $N$ that are perfectly divisible by $q$ (with out leaving a remainder). <br>
**Hint**: $N$ should also be inclusive in the numbers being considered. If N is negative use N == 20

```python
def perfectly_divisible(N, q):
    ...
```

```python
grader.check("q2")
```

Output:
```
Question results:
    Question - 1 result:
        Trying:
            perfectly_divisible(20,2)
        Expecting:
            [0, 2, 4, 6, 8, 10, 12, 14, 16, 18, 20]
        **********************************************************************
        Line 2, in Question 0
        Failed example:
            perfectly_divisible(20,2)
        Expected:
            [0, 2, 4, 6, 8, 10, 12, 14, 16, 18, 20]
        Got nothing

    Question - 2 result:
        Trying:
            perfectly_divisible(10,2)
        Expecting:
            [0, 2, 4, 6, 8, 10]
        **********************************************************************
        Line 2, in Question 1
        Failed example:
            perfectly_divisible(10,2)
        Expected:
            [0, 2, 4, 6, 8, 10]
        Got nothing

    Question - 3 result:
        Trying:
            perfectly_divisible(30,2)
        Expecting:
            [0, 2, 4, 6, 8, 10, 12, 14, 16, 18, 20, 22, 24, 26, 28, 30]
        ********************
```

**Question 3:**  Write a function `flatten_lists` that takes in a list of lists and outputs a sorted list of elements of sublists of the input list (confusing right?) <br>
Example: given `flatten_lists([[2,13,44], [6,7]])` it should return `[2,6,7,13,44]`

```python
def flatten_lists(superlist):
    ...
```

```python
grader.check("q3")
```

Output:
```
Question results:
    Question - 1 result:
        Trying:
            flatten_lists([[2,13,44], [6,7]])
        Expecting:
            [2, 6, 7, 13, 44]
        **********************************************************************
        Line 2, in Question 0
        Failed example:
            flatten_lists([[2,13,44], [6,7]])
        Expected:
            [2, 6, 7, 13, 44]
        Got nothing

    Question - 2 result:
        Trying:
            flatten_lists([[2,13,12], [6,7]])
        Expecting:
            [2, 6, 7, 12, 13]
        **********************************************************************
        Line 2, in Question 1
        Failed example:
            flatten_lists([[2,13,12], [6,7]])
        Expected:
            [2, 6, 7, 12, 13]
        Got nothing

    Question - 3 result:
        Trying:
            flatten_lists([[22,13,4], [67,7]])
        Expecting:
            [4, 7, 13, 22, 67]
        ********************************************************************
```

**Question 4:**  Write a function `longest_word` to return the longest string in a line of text <br>
Example: for text `Hallo people of the world!` should return `people` <br>
If there is a tie, then it should return the `1st` "longest" word encountered

```python
def longest_word(sentence):
    ...
```

```python
grader.check("q4")
```

Output:
```
Question results:
    Question - 1 result:
        Trying:
            longest_word('Hallo people of the world!')
        Expecting:
            'people'
        **********************************************************************
        Line 2, in Question 0
        Failed example:
            longest_word('Hallo people of the world!')
        Expected:
            'people'
        Got nothing

    Question - 2 result:
        Trying:
            longest_word('Welcome to DSA 2020 Arusha Tanzania!')
        Expecting:
            'Tanzania!'
        **********************************************************************
        Line 2, in Question 1
        Failed example:
            longest_word('Welcome to DSA 2020 Arusha Tanzania!')
        Expected:
            'Tanzania!'
        Got nothing
```

**Question 5:**  Write a function `dictionarize` that takes a list as input and returns a dictionary with the elements of the list as the keys and the frequency of the element as the values.

```python
def dictionarize(alist):
    ...
```

```python
grader.check("q5")
```

Output:
```
Question results:
    Question - 1 result:
        Trying:
            dictionalize([1,1,1,1,2,2,2,2,3,3,4,5,5])
        Expecting:
            {1: 4, 2: 4, 3: 2, 4: 1, 5: 2}
        **********************************************************************
        Line 2, in Question 0
        Failed example:
            dictionalize([1,1,1,1,2,2,2,2,3,3,4,5,5])
        Exception raised:
            Traceback (most recent call last):
              File "/usr/lib/python3.8/doctest.py", line 1336, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest Question 0[0]>", line 1, in <module>
                dictionalize([1,1,1,1,2,2,2,2,3,3,4,5,5])
            NameError: name 'dictionalize' is not defined

    Question - 2 result:
        Trying:
            dictionalize([1,1,1,6,8,2,9,3,3,3,4,5,5])
        Expecting:
            {1: 3, 6: 1, 8: 1, 2: 1, 9: 1, 3: 3, 4: 1, 5: 2}
        **************************************************************
```

**Question 6:**  Write a function `greatest_common_divisor` that takes two inputs `a` and `b` and returns the greatest common divisor of the two numbers. E.g. input `(10, 15)` would return `5`

```python
def greatest_common_divisor(a, b):
    ...
```

```python
grader.check("q6")
```

Output:
```
Question results:
    Question - 1 result:
        Trying:
            greatest_common_divisor(10,15)
        Expecting:
            5
        **********************************************************************
        Line 2, in Question 0
        Failed example:
            greatest_common_divisor(10,15)
        Expected:
            5
        Got nothing

    Question - 2 result:
        Trying:
            greatest_common_divisor(20,50)
        Expecting:
            10
        **********************************************************************
        Line 2, in Question 1
        Failed example:
            greatest_common_divisor(20,50)
        Expected:
            10
        Got nothing
```

**Question 7:**  Write a function `get_nearest_farthest` that takes in a point of interest (pt) and a list of points and returns the nearest point and the farthest point from the point of intest.

```python
def get_nearest_farthest(pt, ptlist):
    ...
```

```python
grader.check("q7")
```

Output:
```
Question results:
    Question - 1 result:
        Trying:
            get_nearest_farthest([1.5,0.5], [[1,2],[2,1],[2,3],[4,5],[6,7],[7,6],[6,8],[9,10]])
        Expecting:
            ([2, 1], [9, 10])
        **********************************************************************
        Line 2, in Question 0
        Failed example:
            get_nearest_farthest([1.5,0.5], [[1,2],[2,1],[2,3],[4,5],[6,7],[7,6],[6,8],[9,10]])
        Expected:
            ([2, 1], [9, 10])
        Got nothing

    Question - 2 result:
        Trying:
            get_nearest_farthest([1.5,0.5], [[1,2],[4,1],[2,3],[4,5],[6,7],[7,6],[6,8],[9,1]])
        Expecting:
            ([1, 2], [6, 8])
        **********************************************************************
        Line 2, in Question 1
        Failed example:
            get_nearest_farthest([1.5,0.5], [[1,2],[4,1],[2,3],[4,5],[6,7],[7,6],[6,8],[9,1]])
        Expected:
            ([1, 2], [6, 8])
        Got nothing
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
Question results:
    Question - 1 result:
        Trying:
            get_nearest_farthest([1.5,0.5], [[1,2],[2,1],[2,3],[4,5],[6,7],[7,6],[6,8],[9,10]])
        Expecting:
            ([2, 1], [9, 10])
        **********************************************************************
        Line 2, in Question 0
        Failed example:
            get_nearest_farthest([1.5,0.5], [[1,2],[2,1],[2,3],[4,5],[6,7],[7,6],[6,8],[9,10]])
        Expected:
            ([2, 1], [9, 10])
        Got nothing

    Question - 2 result:
        Trying:
            get_nearest_farthest([1.5,0.5], [[1,2],[4,1],[2,3],[4,5],[6,7],[7,6],[6,8],[9,1]])
        Expecting:
            ([1, 2], [6, 8])
        **********************************************************************
        Line 2, in Question 1
        Failed example:
            get_nearest_farthest([1.5,0.5], [[1,2],[4,1],[2,3],[4,5],[6,7],[7,6],[6,8],[9,1]])
        Expected:
            ([1, 2], [6, 8])
        Got nothing
```

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Output:
```
There was an error generating your LaTeX; showing full error message:
================================================================================
Failed to run "['xelatex', 'notebook.tex', '-quiet']" command:
This is XeTeX, Version 3.14159265-2.6-0.999991 (TeX Live 2019/Debian) (preloaded format=xelatex)
 restricted \write18 enabled.
entering extended mode
(./notebook.tex
LaTeX2e <2020-02-02> patch level 2
L3 programming layer <2020-02-14>
(/usr/share/texlive/texmf-dist/tex/latex/base/article.cls
Document Class: article 2019/12/20 v1.4l Standard LaTeX document class
(/usr/share/texlive/texmf-dist/tex/latex/base/size10.clo))
(/usr/share/texlive/texmf-dist/tex/xelatex/xecjk/xeCJK.sty
(/usr/share/texlive/texmf-dist/tex/latex/l3kernel/expl3.sty
(/usr/share/texlive/texmf-dist/tex/latex/l3backend/l3backend-xdvipdfmx.def))
(/usr/share/texlive/texmf-dist/tex/latex/l3packages/l3keys2e/l3keys2e.sty)
(/usr/share/texlive/texmf-dist/tex/latex/l3packages/xtemplate/xtemplate.sty)
(/usr/share/tex
```

Credit: [Data Science Africa 2022](https://www.datascienceafrica.org/dsa2022arusha/blog/summer-school)

Tests for q1 (points: 6)
```
>>> ypoly(3,3) == 139
True
```
```
>>> ypoly(-2,3) == 0
True
```


Tests for q2 (points: 5)
```
>>> assert_equal(perfectly_divisible(10,2), [0,2,4,6,8,10])
```


Tests for q3 (points: 4)
```
>>> assert_equal(flatten_lists([[2,13,44], [6,7]]), [2, 6, 7, 13, 44])
```


Tests for q4 (points: 5)
```
>>> assert_equal(longest_word("Hello there people"), 'people')
```


Tests for q5 (points: 5)
```
>>> assert_equal(dictionarize([1,2,2,3,4,5]), {1: 1, 2: 2, 3: 1, 4: 1, 5: 1})
```


Tests for q6 (points: 5)
```
>>> assert_equal(greatest_common_divisor(10, 15), 5)
```


Tests for q7 (points: 10)
```
>>> assert_equal(get_nearest_farthest((3, 8), [(9, 3), (8,5), (7,6)]), ((7, 6), (9, 3)))
```
